GABA-ergic inhibition in human hMT+ predicts visuo-spatial intelligence mediated through the frontal cortex

Author response:

The following is the authors’ response to the previous reviews.

Reviewer #1 (Public Review):

Summary:

The study of human intelligence has been the focus of cognitive neuroscience research, and finding some objective behavioral or neural indicators of intelligence has been an ongoing problem for scientists for many years. Melnick et al, 2013 found for the first time that the phenomenon of spatial suppression in motion perception predicts an individual's IQ score. This is because IQ is likely associated with the ability to suppress irrelevant information. In this study, a high-resolution MRS approach was used to test this theory. In this paper, the phenomenon of spatial suppression in motion perception was found to be correlated with the visuo-spatial subtest of gF, while both variables were also correlated with the GABA concentration of MT+ in the human brain. In addition, there was no significant relationship with the excitatory transmitter Glu. At the same time, SI was also associated with MT+ and several frontal cortex FCs.

Strengths:

(1) 7T high-resolution MRS is used.

(2) This study combines the behavioral tests, MRS, and fMRI.

Weaknesses:

Major:

In Melnick (2013) IQ scores were measured by the full set of WAIS-III, including all subtests. However, this study only used visual spatial domain of gF. I wonder why only the visuo-spatial subtest was used not the full WAIS-III? I am wondering whether other subtests were conducted and, if so, please include the results as well to have comprehensive comparisons with Melnick (2013).

We thank the reviewer for pointing this out. The decision was informed by Melnick’s findings which indicated high correlations between Surround suppression (SI) and the Verbal Comprehension, Perceptual Reasoning, Working Memory, and Processing Speed Indexes, with correlation coefficients of 0.69, 0.47, 0.49, and 0.50, respectively. It is well-established that the hMT+ region of the brain is a sensory cortex involved in visual perception processing (3D perception). Furthermore, motion surround suppression (SI), a specific function of hMT+, aligns closely with this region's activities. Given this context, the Perception Reasoning sub-ability was deemed to have the clearest mechanism for further exploration. Consequently, we selected the most representative subtest of Perception Reasoning—the Block Design Test—which primarily assesses 3D visual intelligence.” For further clarification, due to these reasons, we conducted only the visuo-spatial subtest.

Minor:

Comments:

In the first revised version, we addressed the following recommendations in the 'Author response' file titled 'Recommendation for the authors.' It seems our response may not have reached you successfully. We would like to share and expand upon our response here:

(1) Table 1 and Table supplementary 1-3 contain many correlation results. But what are the main points of these values? Which values do the authors want to highlight? Why are only p-values shown with significance symbols in Table supplementary 2??

(1.1) What are the main points of these values?

Thank reviewer for pointing this out. These correlations represent the relationship between behavior task (SI/BDT) and resting-state functional connectivity. It indicates that left hMT+ is involved in the efficient information integration network when it comes to BDT task. In addition, left hMT+’s surround suppression is involved in several hMT+ - frontal connectivity. Furthermore, the overlap regions between two task indicates the underlying mechanism.

(1.2) Which values do the authors want to highlight?

Table 1 and Table Supplementary 1-3 present the preliminary analysis results for Table 2 and Table Supplementary 4-6. So, we generally report all value. Conversely, in the Table 2 and Table Supplementary 4-6, we highlight the value which support our main conclusion.

(1.3) Why are only p-values shown with significance symbols in Table Supplementary 2?

Thank you for pointing this out, it is a mistake. We have revised it and delete the significance symbols.

(2) Line 27, it is unclear to me what is "the canonical theory".

We thank reviewer for pointing this out. We have revised “the canonical theory" to “the prevailing opinion” (line 27)

(3) Throughout the paper, the authors use "MT+", I would suggest using "hMT+" to indicate the human MT complex, and to be consistent with the human fMRI literature.

We thank reviewer for pointing this out. We have revised them.

(4) At the beginning of the results section, I suggest including the total number of subjects. It is confusing what "31/36 in MT+, and 28/36 in V1" means.

We thank reviewer for pointing this out. We have included the total number of subjects in the beginning of result section. (line 110, line 128)

(5) Line 138, "This finding supports the hypothesis that motion perception is associated with neural activity in MT+ area". This sentence is strange because it is a well-established finding in numerous human fMRI papers. I think the authors should be more specific about what this finding implies.

We thank reviewer for pointing this out. We have revised it to:” This finding is in line with prior results, which indicates that motion perception is associated with neural activity in hMT+ area, but not in EVC (primarily in V1)” (lines 156-158)

(6) There are no unit labels for all x- and y-axies in Figure 1. I only see the unit for Conc is mmol per kg wet weight.

We thank reviewer for pointing this out. Figure 1 is a schematic and workflow chart, so labels for x- and y-axes are not needed. I believe this confusion might pertain to Figure 3. In Figures 3a and 3b, the MRS spectrum does not have a standard y-axis unit as it varies based on the individual physical conditions of the scanner; it is widely accepted that no y-axis unit is used. While the x-axis unit is ppm, which indicate the chemical shift of different metabolites. In Figure 3c, the BDT represents IQ scores, which do not have a standard unit. Similarly, in Figures 3d and 3e, the Suppression Index does not have a standard unit.

(7) Although the correlations are not significant in Figure Supplement 2&3, please also include the correlation line, 95% confidence interval, and report the r values and p values (i.e., similar format as in Figure 1C).

We thank reviewer for pointing this out. We have revised them and include the correlation line, 95% confidence interval, r values and p values.

(8) There is no need to separate different correlation figures into Figure Supplementary 1-4. They can be combined into the same figure.

We thank reviewer for the suggestion. However, each correlation figure in the supplementary figures has its own specific topic and conclusion. Please notes that in the revised version, we have added a figure showing the EVC (primarily in V1) MRS scanning ROI as Supplementary Figure 1. Therefore, the figures the reviewer is concerned about are Supplementary Figure 2-5. The correlation figures in Supplementary Figure 2 indicate that GABA in EVC (primarily in V1) does not show any correlation with BDT and SI, illustrating that inhibition in EVC (primarily in V1) is unrelated to both 3D visuo-spatial intelligence and motion suppression processing. The correlations in Supplementary Figure 3 indicate that the excitation mechanism, represented by Glutamate concentration, does not contribute to 3D visuo-spatial intelligence in either hMT+ or EVC (primarily in V1). Supplementary Figure 4 validates our MRS measurements. Supplementary Figure 5 addresses potential concerns regarding the impact of outliers on correlation significance. Even after excluding two “outliers” from Figures 3d and 3e, the correlation results remain stable.

(9) Line 213, as far as I know, the study (Melnick et al., 2013) is a psychophysical study and did not provide evidence that the spatial suppression effect is associated with MT+.

We thank reviewer for pointing this out. It was a mistake to use this reference, and we have revised it accordingly. (line 242)

(10) At the beginning of the results, I suggest providing more details about the motion discrimination tasks and the measurement of the BDT.

We thank reviewer for pointing this out. We have included some brief description of task in the beginning of result section. (lines 116-120)

(11) Please include the absolute duration thresholds of the small and large sizes of all subjects in Figure 1.

We thank reviewer for the suggestion. We have included these results in Figure 3.

(12) Figure 5 is too small. The items in plot a and b can be barely visible.

We thank reviewer for pointing this out. We increase the size and resolution of the Figure.

Reviewer #3 (Public Review):

(1) Throughout the manuscript, hMT+ connectivity with the frontal cortex has been treated as an a priori hypothesis/space. However, there is no such motivation or background literature mentioned in the Introduction. Can the authors clarify the necessity of functional connectivity? In other words, can BOLD activity of hMT+ in the localizer task substitute for functional connectivity between hMT+ and the frontal cortex?

(1.1) Throughout the manuscript, hMT+ connectivity with the frontal cortex has been treated as an a priori hypothesis/space. However, there is no such motivation or background literature mentioned in the Introduction. Can the authors clarify the necessity of functional connectivity?

We thank reviewer for pointing this out. We offered additional motivation and background literature in the introduction: “Frontal cortex is usually recognized as the cognitive core region (Duncan et al., 2000; Gray et al., 2003). Strong connectivity between the cognitive regions suggests a mechanism for large-scale information exchange and integration in the brain (Barbey, 2018; Cole et al., 2012). Therefore, the potential conjunctive coding may overlap with the inhibition and/or excitation mechanism of hMT+. Taken together, we hypothesized that 3D visuo-spatial intelligence (as measured by BDT) might be predicted by the inhibitory and/or excitation mechanisms in hMT+ and the integrative functions connecting hMT+ with frontal cortex (Figure 1a).” (lines 67-74). Additionally, we have included a whole-brain analysis for validation. Functional connectivity reveals the information exchange relationships across regions, enhancing our understanding of how hMT+ and the frontal cortex collaborate when solving visual-spatial intelligence tasks.

(1.2) In other words, can BOLD activity of hMT+ in the localizer task substitute for functional connectivity between hMT+ and the frontal cortex?

We thank the reviewer for this question. The localizer task was used solely for defining the hMT+ MRS scanning area. Functional connectivity was measured using resting-state fMRI. Research has shown that resting-state functional connectivity between the frontal cortex and other ROIs can further reveal the neural mechanisms underlying intelligence tasks (Song et al., 2008).

(2) There is an obvious mismatch between the in-text description and the content of the figure:

"In contrast, there was no correlation between BDT and GABA levels in V1 voxels (figure supplement 1a). Further, we show that SI significantly correlates with GABA levels in hMT+ voxels (r = 0.44, P = 0.01, n = 31, Figure 3d). In contrast, no significant correlation between SI and GABA concentrations in V1 voxels was observed (figure supplement 1b)."

We thank reviewer for pointing this out. We have revised it. The revised version is :” In contrast, there was no correlation between BDT and GABA levels in V1 voxels (figure supplement 2a). Further, we show that SI significantly correlates with GABA levels in hMT+ voxels (r = 0.44, P = 0.01, n = 31, Figure 3d). In contrast, no significant correlation between SI and GABA concentrations in V1 voxels was observed (figure supplement 2b).” (lines 151-156)

(3) The authors' response to my previous round of review indicated that the "V1 ROIs" covered a substantial amount of V3 (32%). Therefore, it would no longer be appropriate to call these "V1 ROIs". I'd suggest renaming them as "Early Visual Cortex (EVC) ROIs" to be more accurate. Can the authors justify why choosing the left hemisphere for visual intelligence task, which is typically believed to be right lateralized?

(3.1) The authors' response to my previous round of review indicated that the "V1 ROIs" covered a substantial amount of V3 (32%). Therefore, it would no longer be appropriate to call these "V1 ROIs". I'd suggest renaming them as "Early Visual Cortex (EVC) ROIs" to be more accurate.

We thank the reviewer for pointing this out. We have revised our description of the MRS scanning ROIs to Early Visual Cortex (EVC). Since the majority of our EVC ROIs are in V1 (around 70%) and almost no V2 was included, we decided to mark the EVC ROIs with the explanation "primarily in V1" for better clarification. This terminology has been widely used to better emphasize the V1-based experimental design.

(3.2) Can the authors justify why choosing the left hemisphere for visual intelligence task, which is typically believed to be right lateralized?

We thank the reviewer for pointing this out. The use of the left MT/V5 as a target was motivated by studies demonstrating that left MT+/V5 TMS is more effective at causing perceptual effects (Tadin et al., 2011). Therefore, we chose to use the left hMT+ as our MRS ROI and maintain consistency across different models' ROIs. Additionally, our results support the notion that the visual intelligence task is right lateralized in the frontal cortex. At the resting-fMRI level, we found that significant ROIs, where functional connectivity is highly correlated with BDT scores, are in the right frontal cortex (Figure 5a, b).

(4) "Small threshold" and "large threshold" are neither standard descriptions, and it is unclear what "small threshold" refers to in the following figure caption. Additionally, the unit (ms) is confusing. Does it refer to timing?

"(f) Peason's correlation showing significant negative correlations between BDT and small threshold."

Thank you for pointing this out; we agree with your suggestion. We have revised the terms “small threshold” and “large threshold” to “duration threshold of small grating” and “duration threshold of large grating”, respectively. The unit (ms) refers to timing. The details are described in the methods section: “The duration was adaptively adjusted in each trial, and duration thresholds were estimated using a staircase procedure. Thresholds for large and small gratings were obtained from a 160-trial block that contained four interleaved 3-down/1-up staircases. For each participant, we computed the correct rate for different stimulus durations separately for each stimulus size. These values were then fitted to a cumulative Gaussian function, and the duration threshold corresponding to the 75% correct point on the psychometric function was estimated for each stimulus size”.

(5) In the response letter, the authors mentioned incorporating the neural efficiency hypothesis in the Introduction, but the revised Introduction does not contain such information.

We thank the reviewer for pointing this out. In our revised version, the second paragraph of the introduction addresses the neural efficiency hypothesis: “The “neuro-efficiency” hypothesis is one explanation for individual differences in gF (Haier et al., 1988). This hypothesis puts forward that the human brain’s ability to suppress irrelevant information leads to more efficient cognitive processing. Correspondingly, using a well-known visual motion paradigm (center-surround antagonism) (Liu et al., 2016; Tadin et al., 2003), Melnick et al found a strong link between suppression index (SI) of motion perception and the scores of the block design test (BDT, a subtest of the Wechsler Adult Intelligence Scale (WAIS), which measures the visuo-spatial component (3D domain) of gF (Melnick et al., 2013). Motion surround suppression (SI), a specific function of human extrastriate cortical region, middle temporal complex (hMT+), aligns closely with this region's activities (Gautama & Van Hulle, 2001). Furthermore, hMT+ is a sensory cortex involved in visual perception processing (3D domain) (Cumming & DeAngelis, 2001). These findings suggest that hMT+ potentially plays a significant role in 3D visuo-spatial intelligence by facilitating the efficient processing of 3D visual information and suppressing irrelevant information. However, more evidence is needed to uncover how the hMT+ functions as a core region for 3D visuo-spatial intelligence.” (lines 51-66)

Recommendations for the authors:

Reviewer #1 (Recommendations for The Authors):

In the Code availability, it states that "this paper does not report original code". It seems weird because at least the code to reproduce the figures from the data should be provided.

Thank you for pointing this out. Almost all figures were created using software such as DPABI, BrainNet, and GraphPad Prism 9.5, which are manually operated and do not require code adjustments. However, for the MRS fitting curve, we can provide our MATLAB code for redrawing the MRS fitting. The code has been uploaded to GitHub.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The prevailing opinion emphasizes fronto-parietal network (FPN) is key in mediating general fluid intelligence (gF). Meanwhile, recent studies show that human MT complex (hMT+), located at the occipito-temporal border and involved in 3D perception processing, also plays a key role in gF. However, the underlying mechanism is not clear, yet. To investigate this issue, our study targets visuo-spatial intelligence, which is considered to have high loading on gF. We use ultra-high field magnetic resonance spectroscopy (MRS) to measure GABA/glutamate concentrations in hMT+ combined with resting-state fMRI functional connectivity (FC), behavioral examinations including hMT+ perception suppression test, and a gF subtest for its visuo-spatial component. Our findings show that both GABA in hMT+ and frontal-hMT+ functional connectivity significantly correlate with the performance of visuo-spatial intelligence. Further, serial mediation model demonstrates that the effect of hMT+ GABA on visuo-spatial gF is fully mediated by the hMT+ frontal FC. Together, our findings highlight the importance in integrating sensory and frontal cortices in mediating the visuospatial component of general fluid intelligence.

Revised version: The prevailing opinion emphasizes fronto-parietal network (FPN) is key in mediating general fluid intelligence (gF). Meanwhile, recent studies show that human MT complex (hMT+), located at the occipito-temporal border and involved in 3D perception processing, also plays a key role in gF. However, the underlying mechanism is not clear, yet. To investigate this issue, our study targets visuo-spatial intelligence, which is considered to have high loading on gF. We use ultra-high field magnetic resonance spectroscopy (MRS) to measure GABA/glutamate concentrations in hMT+ combining resting-state fMRI functional connectivity (FC), behavioral examinations including hMT+ perception suppression test and gF subtest in visuo- spatial component. Our findings show that both GABA in hMT+ and frontal-hMT+ functional connectivity significantly correlate with the performance of visuo-spatial intelligence. Further, serial mediation model demonstrates that the effect of hMT+ GABA on visuo-spatial gF is fully mediated by the hMT+ frontal FC. Together our findings highlight the importance in integrating sensory and frontal cortices in mediating the visuospatial component of general fluid intelligence.

Introduction

Earlier version: General fluid intelligence (gF) is a current problem-solving ability, which shows high inter-individual differences in humans(Cattell & Raymond, 1963). At the beginning of the last century, Spearman(Spearman, 1904) proposed that some general or g factor contributes to our gF. One key component of gF is visuo-spatial intelligence that, usually tested by visual materials, shows high g-loading(Colom et al., 2006; Deary et al., 2010; Jung & Haier, 2007). The exact neural mechanisms of the interplay of visuo spatial intelligence with gF remain yet unclear, though.

Revised version: General fluid intelligence (gF) is a current problem-solving ability, which shows high inter-individual differences in humans(Cattell & Raymond, 1963). At the beginning of the last century, Spearman(Spearman, 1904) proposed that some general or g factor contributes to our gF. One key component of gF is visuo-spatial intelligence, usually tested by visual materials, shows high g-loading(Colom et al., 2006; Deary et al., 2010; Jung & Haier, 2007). The exact neural mechanisms of the interplay of visuo spatial intelligence with gF remain yet unclear, though.

Earlier version: The“neuro-efficiency” hypothesis is one explanation for individual differences in gF (Haier et al., 1988). This hypothesis puts forward that the human brain’s ability to suppress irrelevant information leads to more efficient cognitive processing. Correspondingly, using a well-known visual motion paradigm (center-surround antagonism)(Liu et al., 2016; Tadin et al., 2003), Melnick et al found a strong link between suppression index (SI) of motion perception and the scores of the block design test (BDT, a subtest of the Wechsler Adult Intelligence Scale (WAIS), which measures the visuo-spatial component (3D domain) of gF (Melnick et al., 2013). Motion surround suppression (SI), a specific function of human extrastriate cortical region, middle temporal complex (hMT+), aligns closely with this region’s activities (Gautama & Van Hulle, 2001). Furthermore, hMT+ is a sensory cortex involved in visual perception processing (3D domain) (Cumming & DeAngelis, 2001). These findings suggest that hMT+ potentially plays a significant role in 3D visuo-spatial gF by facilitating the efficient processing of 3D visual information and suppressing irrelevant information. However, more evidence is needed to uncover how the hMT+ functions as a core region for 3D visuo-spatial intelligence. We here hypothesize a key role of hMT+ GABA and its functional connectivity with the frontal cortex in mediating visuospatial gF(Figure 1a).

Revised version: The“neuro-efficiency” hypothesis is one explanation for individual differences in gF (Haier et al., 1988). This hypothesis puts forward that the human brain’s ability to suppress irrelevant information leads to more efficient cognitive processing. Correspondingly, using a well-known visual motion paradigm (center-surround antagonism)(Liu et al., 2016; Tadin et al., 2003), Melnick et al found a strong link between suppression index (SI) of motion perception and the scores of the block design test (BDT, a subtest of the Wechsler Adult Intelligence Scale (WAIS), which measures the visuo-spatial component (3D domain) of gF (Melnick et al., 2013). Motion surround suppression (SI), a specific function of human extrastriate cortical region, middle temporal complex (hMT+), aligns closely with this region’s activities (Gautama & Van Hulle, 2001). Furthermore, hMT+ is a sensory cortex involved in visual perception processing (3D domain) (Cumming & DeAngelis, 2001). These findings suggest that hMT+ potentially plays a significant role in 3D visuo-spatial gF by facilitating the efficient processing of 3D visual information and suppressing irrelevant information. However, more evidence is needed to uncover how the hMT+ functions as a core region for 3D visuo-spatial intelligence.

Added in the revised version: Frontal cortex is usually recognized as the cognitive core region (Duncan et al., 2000; Gray et al., 2003). Strong connectivity between the cognitive regions suggests a mechanism for large-scale information exchange and integration in the brain(Barbey, 2018; Cole et al., 2012). Therefore, the potential conjunctive coding may overlap with the inhibition and/or excitation mechanism of hMT+. Taken together, we hypothesized that 3D visuo-spatial intelligence (as measured by BDT) might be predicted by the inhibitory and/or excitation mechanisms in hMT+ and the integrative functions connecting hMT+ with frontal cortex (Figure 1a).

Earlier version: Hypothesis and experimental design. (a) Schematic of hypothesis. The inhibition mechanism centered on hMT+, including the molecular level: the GABAergic inhibition in hMT+ (green circle), brain connectivity level: hMT+ - frontal functional connectivity (blue circle), and behavior level: hMT+ specific surround suppression of visual motion (red circle), contributes to the visuo-spatial component of gF (3D domain, yellow circle). (b) Schematic of experimental design. Session 1 (rectangle box of short line) was the functional MRI and MRS scanning at resting state. Session 2 (rectangle box of solid line) was another region of MRS acquisition. In the two sessions, the order of MRS scanning regions (hMT+ and V1) was counter balanced across participants. There was a structural MRI scanning before each MRS data acquisition. The interval between the two sessions was used for behavioral measurement (rectangle box of dotted line): block design task (BDT) and psychophysical task - motion discrimination. Sold lines indicate the experiment sequence.

Revised version: Hypothesis and experimental design. (a) Schematic of hypothesis. The inhibition mechanism centered on MT+ GABA, including the molecular level: the GABAergic inhibition in MT+ (green circle), brain connectivity level: hMT+ - frontal functional connectivity (blue circle), and behavior level: hMT+ specific surround suppression of visual motion ( red circle), contributes to the visuo-spatial component of gF (3D domain, yellow circle). (b) Schematic of experimental design. Session 1 (rectangle box of short line) was the functional MRI and MRS scanning at resting state. Session 2 (rectangle box of solid line) was another region of MRS acquisition. In the two sessions, the order of MRS scanning regions (hMT+ and EVC (primarily in V1)) was counter balanced across participants. There was a structure MRI scanning before each MRS data acquisition. The interval between the two sessions was used for behavioral measurement (rectangle box of dotted line): block design task (BDT) and psychophysical task - motion discrimination. Sold lines indicate the experiment sequence.

Earlier version: To investigate our hypothesis, this work conducted multi-level examination including biochemical (Glutamatergic - GABAergic in hMT+), regional-systemic (brain connectivity with hMT+ - based), and behavioral (visual motion function in hMT+) levels to reveal if hMT+ contributes to the 3D visuo-spatial component of gF. We employ ultra-high field (7T) magnetic resonance spectroscopy (MRS) technology to reliably resolve GABA and Glu concentrations(Ende, 2015; Liu et al., 2022; Song et al., 2021). To verify the specificity of hMT+, we used primary visual cortex (V1) - based GABA/Glu as control as it mediates the 2D rather than 3D visual domain (Cumming & DeAngelis, 2001).

Revised version: To investigate our hypothesis, this work conducted multi-level examination including biochemical (Glutamatergic - GABAergic in hMT+), regional-systemic (brain connectivity with hMT+ - based), and behavioral (visual motion function in hMT+) levels to reveal if hMT+ contributes to the 3D visuo-spatial component of gF. We employ ultra-high field (7T) magnetic resonance spectroscopy (MRS) technology to reliably resolve GABA and Glu concentrations(Ende, 2015; Liu et al., 2022; Song et al., 2021). To verify the specificity of hMT+, we used early visual cortex (EVC, primarily in V1) - based GABA/Glu as control as it mediates the 2D rather than 3D visual domain(Cumming & DeAngelis, 2001).

Earlier version: Our findings first demonstrate that GABAergic inhibition mechanisms (but not excitatory Glu) in hMT+ region relate to 3D visuo-spatial ability. Further, analysis of functional brain connectivity at rest reveals that the network (between hMT+ and frontal cortex) relating to hMT+ GABA and perceptual suppression contribute to visuo-spatial gF. Our results provide direct evidence that GABA-ergic inhibitory mechanisms in hMT+ region (a sensory cortex) mediate the multi-level visuo-spatial component (3D domain) of gF thus drawing a direct connection of biochemistry, brain connectivity, and behavior.

Revised version: Our findings first demonstrate that GABAergic inhibition mechanisms (but not excitatory Glu) in hMT+ region relate to 3D visuo-spatial ability. Further, analysis of functional brain connectivity at rest reveals that the network (between MT+ and frontal cortex) relating to MT+ GABA and perceptual suppression contribute the visuo-spatial intelligence. Our results provide direct evidence that inhibitory mechanisms centered on GABA levels in MT+ region (a sensory cortex) mediate multi-level visuo-spatial component (3D domain) of gF thus drawing a direct connection of biochemistry, brain connectivity, and behavior.

Results

Earlier version: To determine whether the function of hMT+ cortex contributes to visuo-spatial component (3D domain) of gF, we adopted the experimental design depicted in Figure 1b. Thirty-six healthy subjects participated in this study. Participants underwent two MRI sessions: the first encompassing resting-state fMRI and magnetic resonance spectra (MRS), and the second solely involving MRS. A 30-minute interval separated these sessions, during which participants performed motion discrimination tasks (using center-surround antagonism stimuli)(Tadin et al., 2003) and the block design test (BDT), which assesses the visuo-spatial ability (3D domain) of gF (Fangmeier et al., 2006). In the motion discrimination tasks, a grating of either large or small size was randomly presented at the center of the screen. The grating drifted either leftward or rightward, and participants were asked to judge the perceived moving direction. While in the Block design test, participants were asked to rebuild the figural pattern within a specified time limit using a set of red and white blocks. Both the volume-of-interests (VOIs) of MRS scanning in the left hMT+ (targeted brain area) and the left primary visual cortex (V1, control brain area) had dimensions of 2×2×2 cm3, with the MRS scanning sequences being randomized across the two MRS sessions. The hMT+ MRS VOIs were demarcated using an anatomical landmark (Dumoulin et al., 2000). For 14 subjects, we also utilized fMRI to functionally pinpoint the hMT+ to validate the placement of the VOI (Figure 2a, b). The V1 MRS VOIs were anatomically defined (Methods). Here, MRS data after extensive quality control (31/36 in hMT+, and 28/36 in V1) were taken for further analysis (Methods).

Revised version: To determine whether the function of hMT+ cortex contributes to visuo-spatial component (3D domain) of gF, we adopted the experimental design depicted in Figure 1b. Thirty-six healthy subjects participated in this study. Participants underwent two MRI sessions: the first encompassing resting-state fMRI and magnetic resonance spectra (MRS), and the second solely involving MRS. A 30-minute interval separated these sessions, during which participants performed motion discrimination tasks (using center-surround antagonism stimuli)(Tadin et al., 2003) and the block design test (BDT), which assesses the visuo-spatial ability (3D domain) of gF(Fangmeier et al., 2006). In the motion discrimination tasks, a grating of either large or small size was randomly presented at the center of the screen. The grating drifted either leftward or rightward, and participants were asked to judge the perceived moving direction. While in the Block design test, participants were asked to rebuild the figural pattern within a specified time limit using a set of red and white blocks. Both the volume-of-interests (VOIs) of MRS scanning in the left hMT+ (targeted brain area) and the left EVC (primarily in V1, control brain area) had dimensions of 2×2×2 cm3, and the MRS scanning sequences were randomized across the two sessions. The hMT+ MRS VOIs were demarcated using an anatomical landmark(Dumoulin et al., 2000). For 14 subjects, we also utilized fMRI to functionally pinpoint the hMT+ to validate the placement of the VOI (Figure 2a, b). The EVC (primarily in V1) MRS VOIs (figure supplement 1) were anatomically defined (Methods). Here, MRS data after extensive quality control (31/36 in hMT+, and 28/36 in EVC (primarily in V1)) were taken for further analysis (Methods).

hMT+ localizer scans and hMT+ MRS VOI placement. (a) Single task block designs. First: a cross fixation on the center of the screen (10s). Second: a moving grating (2°) toward left last 10s. Third: the grating keeps static for 10s. Fourth: the grating moves toward right last 10s. The localizer scans consist of 8 blocks. (b) hMT+ location and MRS VOI placement. The upper template is the horizontal view. The lower templates from left to right are coronal and sagittal views. The warm color indicates the overlap of fMRI activation of hMT+ across 14 subjects, the cold color bar indicates the overlap of MRS VOIs across all subjects.

Added in the revised version: The left hemisphere EVC (primarily in V1) MRS scanning ROI. The color bar indicates the overlap of the MRS VOIs across all subjects.

Earlier version: GABA and Glu concentrations in hMT+ and V1 and their relation to SI and BDT

Revised version: GABA and Glu concentrations in hMT+ and EVC (primarily in V1) and their relation to SI and BDT

Earlier version: An example of a MRS voxel located in hMT+ is shown in Figure 3a. LCModel fittings for GABA spectra from all subjects in hMT+ (n = 31) and V1 (n = 28) are illustrated in Figure 3b (color scale presents the BDT scores). We discerned a significant association between the inter-subjects’ BDT scores and the GABA levels in hMT+ voxels, but not in V1 voxels. Quantitative analysis displayed that BDT significantly correlates with GABA concentrations in hMT+ voxels (r = 0.39, P = 0.03, n = 31, Figure 3c). After using partial correlation to control for the effect of age, the relationship remains significant (rpartial = 0.43 P = 0.02, 1 participant excluded due to the age greater than mean + 2.5SD). In contrast, there was no correlation between BDT and GABA levels in V1 voxels (figure supplement 1a). Further, we show that SI significantly correlates with GABA levels in hMT+ voxels (r = 0.44, P = 0.01, n = 31, Figure 3d). In contrast, no significant correlation between SI and GABA concentrations in V1 voxels was observed (figure supplement 1b). This finding is in line with prior results, which indicates that motion perception is associated with neural activity in hMT+ area, but not in V1 (Schallmo et al., 2018). LCModel fittings for Glu spectra from all subjects in hMT+ (n = 31) and V1 (n = 28) voxels are presented in figure supplement 2a.

Revised version: An example of a MRS voxel located in hMT+ is shown in Figure 3a. LCModel fittings for GABA spectra from all subjects in hMT+ (n = 31) and EVC (primarily in V1) (n = 28) are illustrated in Figure 3b (color scale presents the BDT scores). We discerned a significant association between the inter-subjects’ BDT scores and the GABA levels in hMT+ voxels, but not in EVC (primarily in V1) voxels. Quantitative analysis displayed that BDT significantly correlates with GABA concentrations in hMT+ voxels (r = 0.39, P = 0.03, n = 31, Figure 3c). After using partial correlation to control for the effect of age, the relationship remains significant (rpartial = 0.426 P = 0.02, 1 participant excluded due to the age greater than mean + 2.5SD) In contrast, there was no obvious correlation between BDT and GABA levels in EVC (primarily in V1) voxels (figure supplement 2a). We show that SI significantly correlates with GABA levels in hMT+ voxels (r = 0.44, P = 0.01, n = 31, Figure 3d). In contrast, no significant correlation between SI and GABA concentrations in EVC (primarily in V1) voxels was observed (figure supplement 2b). This finding is in line with prior results, which indicates that motion perception is associated with neural activity in hMT+ area, but not in EVC (primarily in V1)(Schallmo et al., 2018). LCModel fittings for Glu spectra from all subjects in hMT+ (n = 31) and EVC (primarily in V1) (n = 28) voxels are presented in figure supplement 3a.

Earlier version: MRS spectra and the relationships between GABA levels and SI / BDT. (a) Example spectrum from the hMT+ voxel of one participant. The first line is the LCModel fitting result of all metabolites, and the following lines show the Glu and GABA spectra fitting with LCModel, and then the baseline. (b) Individual participants fitted GABA MRS spectra from the hMT+ (top) and V1 (bottom) voxels from baseline measurement. The colors of the GABA spectra represent the individual differences of BDT. The color bar represents the scores of BDT. (c) and (d) Pearson’s correlations showing significant positive correlations between hMT+ GABA and BDT scores (c), between hMT+ GABA and SI (d). GABA and Glu concentrations (Conc.) are absolute, with units of mmol per kg wet weight (Methods). (e) Peason’s correlation showing significant positive correlations between SI and BDT. (f) Peason’s correlation showing significant negative correlations between BDT and small threshold. (g) No correlation between BDT and large threshold. The ribbon between dotted lines represents the 95% confidence interval, and the black regression line represents the Pearson’s correlation coefficient (r).

Revised version: MRS spectra and the relationships between GABA levels and SI / BDT. (a) Example spectrum from the hMT+ voxel of one participant. The first line is the LCModel fitting result of all metabolites, and the following lines show the Glu and GABA spectra fitting with LCModel, and then the baseline. (b) Individual participants fitted GABA MRS spectra from the hMT+ (top) and EVC (primarily in V1) (bottom) voxels from baseline measurement. The colors of the GABA spectra represent the individual differences of BDT. The color bar represents the scores of BDT. (c) and (d) Pearson’s correlations showing significant positive correlations between hMT+ GABA and BDT scores (c), between hMT+ GABA and SI (d). (e) Peason’s correlation showing significant positive correlations between SI and BDT. (f) Peason’s correlation showing significant negative correlations between BDT and duration threshold of small grating. (g) No correlation between BDT and duration threshold of large grating. The ribbon between dotted lines represents the 95% confidence interval, and the black regression line represents the Pearson’s correlation coefficient (r). GABA and Glu concentrations (Conc.) are absolute, with units of mmol per kg wet weight (Methods).

Added in the revised version: Relationships between GABA concentration in EVC (primarily in V1) and BDT/SI. (a) There is no significant relationship between BDT and GABA concentrations in EVC (primarily in V1) region. (b) There is also no obvious correlation between SI and GABA concentrations in EVC (primarily in V1) region. GABA concentration (Conc.) is absolute, with units of mmol per kg wet weight (Methods).

Added in the revised version: Individual Glu MRS spectra from hMT+ / EVC (primarily in V1) regions and relationships between BDT and Glu concentrations in hMT+ / EVC (primarily in V1) regions.

Added in the revised version: (a) Individual participant fitted Glutamate MRS spectra from the hMT+ (top, n = 31) and EVC (primarily in V1) (bottom, n = 28) voxels from baseline measurement. The colors of the Glutamate spectra represent the individual differences of BDT. The color bar represents the scores of BDT. (b) There is no significant relationship between BDT and Glu concentrations in hMT+ region. c There is also no significant correlation between BDT and Glu concentrations in EVC (primarily in V1) region. Glutamate concentration (Conc.) is absolute, with units of mmol per kg wet weight (Methods).

Added in the revised version: Correlations between GABA and Glu concentrations in hMT+ and EVC (primarily in V1) regions. (a) There is significant correlation between GABA and Glu concentrations in hMT+ region. (b) The levels of GABA also significantly correlate with Glu in EVC (primarily in V1) region. GABA and Glu concentrations (Conc.) are absolute, with units of mmol per kg wet weight (Methods).

Added in the revised version: Two linear correlations. (a) Significant positive correlation between suppression index and BDT scores (took out two outlines). (b) BDT also significantly correlates with GABA concentration in hMT+ region (without two outliers, having the similar result shown in Figure 3d).

Earlier version: Unlike in the case of GABA, no significant correlations between BDT and Glu levels were found in both hMT+ and V1 voxels (figure supplement 2b, c). While, as expected (Song et al., 2021), we observed significant positive correlations between GABA and Glu concentrations in both hMT+ (r = 0.62, P < 0.001, n = 31) and V1 voxels (r = 0.56, P = 0.002, n = 28) (figure supplement 3a, b). Additionally, significant correlations between SI and BDT, small threshold and SI were discerned (r = 0.59, P = 0.0002, n = 34, figure 3e, rpartial = 0.67, P < 0.001, 1 participant excluded due to the age greater than mean + 2.5SD; r = -0.43, P = 0.016, rpartial = 0.44, P = 0.014). While there was no significant correlation between large threshold and SI, corroborating previous conclusions (Melnick et al., 2013). Two outliers evident in Figure 3e were excluded, with consistent results depicted in figure supplement 4a. Further, two outliers evident in Figure 3d were excluded, with consistent results depicted in figure supplement 4b.

Revised version: Unlike in the case of GABA, no significant correlations between BDT and Glu levels were found in both hMT+ and EVC (primarily in V1) voxels (figure supplement 3b, c). While, as expected(Song et al., 2021), we observed significant positive correlations between GABA and Glu concentrations in both hMT+ (r = 0.62, P = 0.0002, n = 31) and EVC (primarily in V1) voxels (r = 0.56, P = 0.002, n = 28) (figure supplement 4a, b). Additionally, significant correlations between SI and BDT, duration threshold of small grating and BDT was discerned (r = 0.59, P = 0.0002, n = 34, Figure 3e, rpartial = 0.67, P < 0.001, 1 participant excluded due to the age greater than mean + 2.5SD; r = - 0.43, P = 0.016, rpartial = 0.44, P = 0.014, Figure 3f). While there was no significant correlation between duration threshold of large grating and BDT (Figure 3g), corroborating previous conclusions(Melnick et al., 2013). Two outliers evident in Figure 3e were excluded, with consistent results depicted in figure supplement 5a. Further, two outliers evident in Figure 3d were excluded, with consistent results depicted in figure supplement 5b.

MT - frontal FC relates to SI and BDT

Earlier version: We next took the left hMT+ as the seed region and separately measured interregional FCs between the seed region and each voxel in the frontal regions (a priori search space). These measurements were correlated with the performance in 3D visuospatial ability (BDT) to identify FCs with significant correlations. Results from connectivity-BDT analysis are summarized in Table 1 and shown in Figure 4a. We found that brain regions with high FC strength to the seed region (left hMT+) significantly correlated with BDT scores, these included mainly regions situated within the canonical cognitive cores of FPN (Brodmann areas (BAs) 6, 9, 10, 46, 47)(Assem et al., 2020; Deary et al., 2010; Duncan et al., 2020; Duncan et al., 2000; Gray et al., 2003; Jung & Haier, 2007). Across the whole brain search, the similar FCs (between hMT+ and frontal cognitive cores) still showed significant correlations with BDT scores (Table supplement 1) (also shown in figure supplement 5a). Additionally, we identified certain parietal regions (BAs 7, 39, 40) with significant correlations of their functional connectivity to the left hMT+ and the BDT scores (Table supplement 1) (also shown in figure supplement 5a). These significant connections between hMT+ and FPN system suggest that left hMT+ is involved in the efficient information integration network mediating the visuo spatial component of gF.

Revised version: We next took the left hMT+ as the seed region and separately measured interregional FCs between the seed region and each voxel in the frontal regions (a priori search space). These measurements were correlated with performance in 3D visuospatial ability (BDT) to identify FCs with significant correlations. Results from connectivity-BDT analysis are summarized in Table 1 and shown in Figure 4a. We found that brain regions with FC strength to the seed region (left hMT+) significantly correlated with BDT scores were situated within the canonical cognitive cores of FPN (Brodmann areas (BAs) 6, 9, 10, 46, 47)(Assem et al., 2020; Deary et al., 2010; Duncan et al., 2020; Duncan et al., 2000; Gray et al., 2003; Jung & Haier, 2007). Across the whole brain search, the similar FCs (between hMT+ and frontal cognitive cores) still showed significant correlations with BDT scores (Table supplement 1) (also shown in figure supplement 6a). Additionally, we identified certain parietal regions (BAs 7, 39, 40) with significant correlations between their connectivity to the left hMT+ and the BDT scores (Table supplement 1) (also shown in figure supplement 6a). These significant connections between hMT+ and FPN system suggest that left hMT+ is involved in the efficient information integration network mediating the visuo spatial component of gF.

Significant FCs from connectivity-behavior analyses in a priori search space. The seed region is the left hMT+. The significant FCs are obtained from a priori space (frontal cortex). (a) The significant FCs obtained from connectivity-BDT analysis. Single voxel threshold P < 0.005, adjacent size ≥ 23 (AlphaSim correcting, Methods). (b) The significant FCs obtained from connectivity-SI analysis. Single voxel threshold P < 0.005, adjacent size ≥ 22 (AlphaSim correcting, Methods). Positive correlations are shown in warm colors, while negative correlations are shown in cold colors. The paired deep green, deep pink, light green circles on (a) and (b) indicate the overlap regions in left BA6, right BA6, right BA46(DLPFC) between Connectivity - BDT analysis and Connectivity – SI analysis, respectively.

Added in the revised version: Significant FCs searched from connectivity-behavior (BDT/SI) analyses in the whole brain. The seed region is the left hMT+. The significant FCs are obtained from the entire brain search, single voxel threshold P < 0.01, adjacent size ≥ 37 voxels (AlphaSim correcting, Methods). Positive correlations are shown in warm colors, while, negative correlations are shown in cold colors. (a) the significant FCs obtained from connectivity-BDT analysis. (b) the significant FCs obtained from connectivity-SI analysis.

Added in the revised version: Significant FCs searched from connectivity-behavior (BDT) analyses in the frontal (a) and whole brain (b). The seed region is the left V1. (a) The significant FCs are obtained from the frontal search, single voxel threshold P < 0.005, adjacent size ≥ 23 (AlphaSim correcting, same methods for hMT+). (b) The significant FCs are obtained from the entire brain search, single voxel threshold P < 0.01, adjacent size ≥ 37 voxels (AlphaSim correcting, same methods for hMT+). Only positive correlations were detected, shown in warm colors.

Earlier version: To address the question whether spatial suppression plays a role, we correlated hMT+ - based global FCs with SI. Though spatial suppression during motion perception (quantified by SI) is considered to be the function of area hMT (Gautama & Van Hulle, 2001), the top-down modulation from the frontal cortex can increase surround suppression (Liu et al., 2016). Our functional connectivity-SI analysis in the frontal regions (a priori search space) displayed 3 brain regions in which FCs strength significantly correlated with SI: right BA4/6, left BA6, and right BA46 (summarized in Table supplement 2, and shown in Figure 4b). Across the whole brain search, we identified total 7 brain regions in which FCs strength significantly correlated with SI, and 3 of these were in the frontal cortex. This is consistent with the results obtained by the functional connectivity-SI analysis in a priori search space (frontal cortex) (Table supplement 3 and figure supplement 5b).

Revised version: To address the question whether spatial suppression plays a role, we correlated hMT+ - based global FCs with SI. Though spatial suppression during motion perception (quantified by SI) is considered to be the function of area hMT+(Gautama & Van Hulle, 2001), the top-down modulation from the frontal cortex can increase surround suppression(Liu et al., 2016). Our functional connectivity-SI analysis in the frontal regions (a priori search space) displayed 3 brain regions in which FCs strength significantly correlated with SI: right BA4/6, left BA6, and right BA46 (summarized in Table supplement 2, and shown in Figure 4b). Across the whole brain search, we identified total 7 brain regions in which FCs strength significantly correlated with SI, and 3 of these were in the frontal cortex. This is consistent with the results obtained by the functional connectivity-SI analysis in a priori search space (frontal cortex) (Table supplement 3 and figure supplement 6b).

Added in the revised version: We also did the V1 functional connectivity-BDT correlations as control analysis (figure supplement 7). Only positive correlations in the frontal area were detected, suggesting that in the 3D visuo-spatial intelligence task, V1 plays a role in feedforward information processing. However, hMT+, which showed specific negative correlations in the frontal, be suggested involving in the inhibition mechanism. These results further emphasize the de-redundancy function of hMT+ in 3D visuo-spatial intelligence.

Earlier version: Local hMT+ GABA acts on SI and BDT via global MT-frontal connectivity

Revised version: Local hMT+ GABA acts on SI and BDT via global hMT-frontal connectivity

To determine whether local neurotransmitter levels (such as GABA and Glu) in the hMT+ region mediate the broader 3D visuo-spatial ability of BDT, which as a component of gF, is linked to the frontal cortex(Fangmeier et al., 2006), we correlated the significant FCs of hMT – frontal in Figure 4a (also shown in Table 1) with the GABA and Glu levels in hMT+ region. The results revealed that only two FCs significantly correlated with inhibitory GABA levels in hMT+: 1) the FC of left hMT+- right BA 46 (significantly negative correlation, r = - 0.56, P = 0.02, n = 29, FDR correction, Figure 5a left); 2) the FC of left hMT+ - right BA 6 (significantly positive correlation, r = 0.69, P = 0.002, n = 29, FDR correction, Figure 5b left) (also shown in Table 2). There were no significant correlations between these FCs and the excitatory Glu levels in hMT+ (Table 2). Across the whole brain search, we obtained the same two hMT+ - frontal FCs significantly correlating with both hMT+ GABA levels and BDT (Table supplement 4), this is consistent with the results in the priori search space (frontal cortex) (Table 2). We then correlated the significant FCs in Figure 5b (also in Table supplement 2) with GABA and Glu concentrations in hMT+ and found that almost all the correlations are significant except one (between the FC of left hMT+ - right BA46 and the Glu levels in hMT+) (Table supplement 5). Among the three FCs, the clusters of two FCs have substantial voxel overlap with the FCs we found by the connectivity-BDT analysis (Figure 5a, b). Across the whole brain search, there were total 7 brain regions in which FCs strength were significantly correlated with SI, all the 7 FCs significantly correlated the hMT+ GABA levels, while, no FC had significant correlation with the hMT+ Glu levels (Table supplement 6).

Earlier version: Local hMT+ GABA acts on SI and BDT via global MT-frontal connectivity. (a) Significant negative correlation between the FC of left hMT+ - right DLPFC (BA46) and hMT+ GABA (FDR correction). (b) Significant positive correlation between the FC of left hMT+ - right (pre) motor cortex (BA4/6) and hMT+ GABA (FDR correction). In (a) and (b), left: the significant FCs obtained from connectivity-BDT analysis; right: the significant FCs obtained from connectivity-SI analysis. (c) Significant pathways: hMT+ GABA→ FC (left hMT+ - right BA46, negative correlation) → SI (negative correlation) → BDT (positive correlation). This pathway can explain 34% of the variance in BDT. (d) Significant pathways: hMT+ GABA → FC (left hMT+ - right BA6, positive correlation) → BDT (positive correlation). The bolded lines represent the hypothesized mediation effect. The dotted lines represent alternative pathways. *: P < 0.05; **: P < 0.01; ***: P < 0.001.

Revised version: Local hMT+ GABA acts on SI and BDT via global hMT-frontal connectivity. (a) Significant negative correlation between the FC of left hMT+ - right DLPFC (BA46) and hMT+ GABA (FDR correction). (b) Significant positive correlation between the FC of left hMT+ - right (pre) motor cortex (BA4/6) and hMT+ GABA (FDR correction). In (a) and (b), left: the significant FCs obtained from connectivity-BDT analysis; right: the significant FCs obtained from connectivity-SI analysis. (c) Significant pathways: hMT+ GABA→ FC (left hMT+ - right BA46, negative correlation) → SI (negative correlation) → BDT (positive correlation). This pathway can explain 34% of the variance in BDT. (d) Significant pathways: hMT+ GABA → FC (left hMT+ - right BA6, positive correlation) → BDT (positive correlation).The bolded lines represent the hypothesized mediation effect. The dotted lines represent alternative pathways. *: P < 0.05; **: P < 0.01; ***: P < 0.001.

Added in the revised version: The alternative serial mediation models from local hMT+ GABA to global performance of BDT. The pathway that hMT+ GABA was predicted to be associated with SI, followed by the FC of hMT+ - BA 46, and then BDT, did not yield the chained mediation effects on BDT. *: P < 0.05; **: P < 0.01; ***: P < 0.001.

Added in the revised version: Correlation matrix shows correlations between all measured in our biochemistry and behaviour level. *: P < 0.05, **: P < 0.01, ***: P < 0.001.

Earlier version: Taken together, our results displayed that the overlap FCs from the analyses of connectivity - behavior (BDT and SI) -GABA are the hMT+ - BA 46 and hMT+ - BA 6 (Figure 5a, b). These results suggest that the coupling of FCs of hMT+ - frontal regions (BA 46 and BA 6) with local hMT+ GABA provides the neural basis for both the simple motion perception (quantified by SI) and the complex 3D visuo-spatial ability (quantified by BDT).

Revised version: Taken together, our results displayed that the overlap FCs from the analyses of connectivity - behavior (BDT and SI) -GABA are the hMT+ - BA 46 and hMT+ - BA 6 (Figure 5a, b). These results suggest that the coupling of FCs of hMT+ - frontal regions (BA 46 and BA 6) coupling with local hMT+ GABA provides the neural basis for both the simple motion perception (quantified by SI) and the complex 3D visuo- spatial ability (quantified by BDT).

Earlier version: In order to fully investigate the potential roles of the multiples variables contributing to BDT scores, serial mediation analyses (Hayes, 2013) were applied to both MR and behavioral data. Following our hypothesis, the independent variable (X) is hMT+ GABA, the dependent variable (Y) is BDT scores, the covariate is the age, and the mediators are FC (M1) and SI (M2). We used the overlap clusters from the analyses of connectivity-BDT-GABA and connectivity-SI-GABA to compute the FC of hMT+ - BA46, 1 participant was excluded due to his age greater than mean + 2.5SD. The serial mediation model is shown in Figure 5c. GABA levels in hMT+ significantly negatively correlated with the FC of hMT+ - BA46 (β = - 0.32, P < 0.001), which in turn significantly negatively correlated with SI (β = - 0.19, P = 0.035), and consequently, significantly positively correlated with BDT (β = 38.5, P = 0.009). Critically, bootstrapped analyses revealed that our hypothesized indirect effect (i.e., hMT+ GABA → FC of hMT+ - BA46 → SI → BDT) was significant (β = 2.28, SE = 1.54, 95% CI = [0.03, 5.94]). The model accounted for 34% of the variance in BDT. However, when considering the MT-BA6 FC as the mediator M1, the serial model does not show a significant indirect effect. Consequently, we explored a mediation model, which revealed that the MT-BA6 FC totally mediates the relationship between GABA and BDS. (Figure 5d). For sensitivity purposes, we tested the alternative models, in which the order of the mediators was reversed. The pathway that hMT+ GABA was predicted to be associated with SI, followed by the FC of hMT+ - BA 46, and then BDT, did not yield the chained mediation effects on BDT (figure supplement 6).

Revised version: In order to fully investigate the potential roles of the multiples variables contributing to BDT scores, serial mediation analyses(Hayes, 2013) were applied to both the MR and behavioral data. Following our hypothesis, the independent variable (X) is hMT+ GABA, the dependent variable (Y) is BDT scores, the covariate is the age, and the mediators are FC (M1) and SI (M2). We used the overlap clusters from the analyses of connectivity-BDT-GABA and connectivity-SI-GABA to compute the FC of hMT+ - BA46, 1 participant was excluded due to his age greater than mean + 2.5SD. The serial mediation model is shown in Figure 5c. GABA levels in hMT+ significantly negatively correlated with the FC of hMT+ - BA46 (β = - 0.32, P = 0.0009), which in turn significantly negatively correlated with SI (β = - 0.19, P = 0.035), and consequently, significantly positively correlated with BDT (β = 38.5, P = 0.009). Critically, bootstrapped analyses revealed that our hypothesized indirect effect (i.e., hMT+ GABA → FC of hMT+ - BA46 → SI → BDT) was significant (β = 2.28, SE = 1.54, 95% CI = [0.03, 5.94]). The model accounted for 34% of the variance in BDT. However, when considering the hMT-BA6 FC as the mediator M1, the serial model does not show a significant indirect effect. Consequently, we explored a mediation model, which revealed that the hMT-BA6 FC totally mediates the relationship between GABA and BDS. (Figure 5d). For sensitivity purposes, we tested the alternative models, in which the order of the mediators was reversed. The pathway that hMT+ GABA was predicted to be associated with SI, followed by the FC of hMT+ - BA 46, and then BDT, did not yield the chained mediation effects on BDT (figure supplement 8).

Earlier version: To summarize (shown in Figure 6), the results from the serial mediation analyses are consistent with our hypothesis. That is, higher GABAergic inhibition in hMT+ relates to stronger negative FC between hMT+ and BA46, leading to enhanced ability for surround suppression (filtering out irrelevant information(Tadin, 2015), which ultimately results in more efficient visual 3D processing as key component of gF (the higher BDT scores).

Revised version: To summarize (shown in Figure 6), the results from the serial mediation analyses are consistent with our hypothesis. That is, higher GABAergic inhibition in hMT+ relates to stronger negative FC between hMT+ and BA46, leading to enhanced ability for surround suppression (filtering out irrelevant information(Tadin, 2015), which ultimately resulting in more efficient visual 3D processing as key component of gF (the higher BDT scores).

Sketch depicting the multi-level inhibitory mechanisms centered on hMT+ GABA contributing to visuo-spatial intelligence. Inhibitory GABA in hMT+ (a sensory cortex, shown in green circle), coupling with the functional connectivity between hMT+ and BA46 (cognitive control core, shown in purple circle), and mediated by motion surround suppression (shown in red circle), contributes to visuo-spatial intelligence (BDT, 3D domain, shown in red and white building blocks). In this sketch, the two-colored parallel arrows show the negative FC between hMT+ and BA46, the colored arrows below the green circle display the inhibition mechanisms centered on hMT+ GABA (2), filtered the irrelevant information in (1) and focused on the efficient visuo-spatial information (3). Black long arrows display the direction of information flow: from input information (1) to visuo-spatial intelligence (4).

Discussion

Earlier version: Here, we provide evidence that hMT+ inhibitory mechanisms mediate processing in the visuo-spatial component (3D domain) of gF on multiple levels, that is, from molecular over brain connectivity to behavior. First, this study found that higher hMT+ inhibitory GABA levels (but not excitatory Glu) relate to FC between hMT+ and BA 46 that contribute to both SI and 3D visuo-spatial intelligence (BDT). Our serial mediation analyses indicate that the inhibitory mechanisms related to hMT+ and its GABA levels (but not Glu), FCs of hMT+ - BA46 coupling with hMT+ inhibitory GABA (but not excitatory Glu), and behavior (SI indexing perceptual suppression in hMT+) predict the inter-subject variance in the 3D visuo-spatial intelligence (BDT) (Figure 5c). Second, we demonstrate discrete GABAergic inhibition mechanisms in hMT+ that mediate the strong FCs between hMT+ - frontal regions (BA46 and BA6): significant negative correlation with the FC of hMT+ - BA 46 (Figure 5a), whereas there is significant positive correlation with the FC of hMT+ - BA 6 (Figure 5b). This indicates that different frontal regions, DLPFC (BA 46) and premotor cortex (BA6), contribute uniquely to gF through hMT+ - based inhibitory mechanisms.

Revised version: Here, we provide evidence that hMT+ inhibitory mechanisms mediate processing in the visuo-spatial component (3D domain) of gF on multiple levels, that is, from molecular over brain connectivity to behavior. First, this study found that higher hMT+ inhibitory GABA levels (but not excitatory Glu) relate to FC between hMT+ and BA 46 that contribute to both SI and BDT. Our serial mediation analyses indicate that the inhibitory mechanisms related to hMT+ and its GABA levels in hMT+ (but not Glu), FCs of hMT+ - BA46 coupling with hMT+ inhibitory GABA (but not excitatory Glu), and behavior (SI indexing perceptual suppression in hMT+) predict the inter-subject variance in the 3D gF task (BDT) (Figure 5c). Second, we demonstrate discrete GABAergic inhibition mechanisms in hMT+ that mediate the strong FCs between hMT+ - frontal regions (BA46 and BA6): significant negative correlation with the FC of hMT+ - BA 46 (Figure 5a), whereas there is significant positive correlation with the FC of hMT+ - BA 6 (Figure 5b). This indicates that different frontal regions, DLPFC (BA 46) and premotor cortex (BA6), contribute uniquely to gF through hMT+ - based inhibitory mechanisms.

Earlier version: The goal of our research is to reveal that the inhibitory (not excitatory) mechanism in hMT+ contributes to multi-level processing in 3D visuo-spatial ability (BDT). Monkey electrophysiological experiments revealed that selective attention gates the visual cortex, including area MT, effectively suppressing the irrelevant information(Everling et al., 2002; Treue & Maunsell, 1996). These findings align with the “neural efficiency” hypothesis of intelligence (Haier et al., 1988), which puts forward the human brain’s ability to suppress the repetition of information. Neural suppression is associated with the balance between excitation and inhibition (EIB), usually represented by covariation between Glutamate and GABA(Ozeki et al., 2009). Here, this study exploited the high spectral resolution afforded by ultrahigh field (7T) MRS to reliably resolve GABA measurement, to adequately discriminate the glutamate and glutamine signals, and to resolve the high accuracy Glu measurement(Ende, 2015). This work implemented the MRS scanning in hMT+ (3D visual domain) and V1 (2D visual domain) regions and found that hMT+ inhibitory GABA (but not excitatory Glu) significantly correlated with 3D visuo-spatial intelligence, i.e., the higher GABA levels in hMT+ (rather than excitatory Glu) relate to higher visual 3D processing (Figure 3c). We searched the global hMT+ - based FCs with the connectivity-BDT analyses (in priori search space and whole brain search to valid), and then, correlated these significant FCs with the GABA and Glu concentrations in hMT+ . We found two FCs (MT+ - BA46 and hMT+ - BA6) significantly correlating with hMT+ inhibitory GABA (whereas no FC significantly correlated with hMT+ excitatory Glu). Accordingly, our results emphasize the importance of hMT+ inhibitory GABA (but not excitatory Glu) in processing the 3D visual-spatial intelligence (BDT).

Revised version: The goal of our research is to reveal that the inhibitory (not excitatory) mechanism in hMT+ contributes to multi-level processing in 3D visuo-spatial ability (BDT). Monkey electrophysiological experiments revealed that selective attention gates the visual cortex, including area MT, effectively suppressing the irrelevant information(Everling et al., 2002; Treue & Maunsell, 1996). These findings align with the “neural efficiency” hypothesis of intelligence(Haier et al., 1988), which puts forward the human brain’s ability to suppress the repetition of information. Neural suppression is associated with the balance between excitation and inhibition (EIB), usually represented by covariation between Glutamate and GABA(Ozeki et al., 2009). Here, this study exploited the high spectral resolution afforded by ultrahigh field (7T) MRS to reliably resolve GABA measurement, to adequately discriminate the glutamate and glutamine signals, and to resolve the high accuracy Glu measurement(Ende, 2015). This work implemented the MRS scanning in hMT+ (3D visual domain) and EVC (primarily in V1) (2D visual domain) regions and found that hMT+ inhibitory GABA (but not excitatory Glu) significantly correlated with BDT, i.e., the higher GABA levels in hMT+ (rather than excitatory Glu) relate to higher visual 3D processing (BDT) (Figure 3c). Basically, this study contains the data of SI, BDT, GABA in MT+ and EVC (primarily in V1), Glu in MT+ and EVC (primarily in V1)-all 6 measurements. we made a correlation matrix to reporting all values in figure supplementary 9.

Added in the revised version: We searched the global hMT+ - based FCs with the connectivity-BDT analyses (in priori search space and whole brain search to valid), and then, correlated these significant FCs with the GABA and Glu concentrations in hMT+. We found two FCs (hMT+ - BA46 and hMT+ - BA6) significantly correlating with hMT+ inhibitory GABA (whereas no FC significantly correlated with hMT+ excitatory Glu). Accordingly, our results emphasize the importance of hMT+ inhibitory GABA (but not excitatory Glu) in processing the 3D visual-spatial intelligence (BDT).

Earlier version: Our recent human study(Song et al., 2021) and other study’s animal experiments(Ozeki et al., 2009; Sato et al., 2016) demonstrated that the conjoint action of inhibition (GABA) and excitation (Glu) underlies visual spatial suppression. In this work, our novel data show the chained mediation effects from local hMT+ GABA to more global 3D visuo-spatial intelligence: hMT+ GABA → FC (MT+ and BA46) → SI → BDT. Thereby, our data indicate that inhibitory mechanisms in hMT+, from the biochemical level of GABA over FC to the behavioral level, can predict the inter- subject variance in the 3D visuo-spatial intelligence (BDT) (Figure 5c).

Revised version: Our recent human study(Song et al., 2021) and other study’s animal experiments(Ozeki et al., 2009; Sato et al., 2016) demonstrated that the conjoint action of inhibition (GABA) and excitation (Glu) underlies visual spatial suppression. In this work, our novel data show the chained mediation effects from local hMT+ GABA to more global BDT: hMT+ GABA → FC (hMT+ and BA46) → SI → BDT. Thereby, our data indicate that inhibitory mechanisms in hMT+, from the biochemical level of GABA over FC to the behavioral level, can predict the inter-subject variance in the 3D gF task (BDT) (Figure 5c).

Earlier version: Another interesting finding reveal that GABAergic inhibition in hMT+ coupling with distinct functional connectivity patterns between BA 46-MT+ and BA6-MT+. A previous human fMRI experiment found that the positive and negative correlations between 3D visuo-spatial intelligence and the activation of frontal regions appeared at different reasoning phases (validation or integration phases during reasoning)(Fangmeier et al., 2006). On the one hand, a monkey electrophysiological experiment reported the delayed modulation from PFC (especially in DLPFC (BA 46))) to area MT during a visual motion task(Zaksas & Pasternak, 2006). Computational models converged with empirical data of awake monkey experiments slowing temporal modulation from PFC to MT/MST(Donner et al., 2009; Siegel et al., 2015; Wang, 2002; Wimmer et al., 2015). On the other hand, human MEG studies(Donner et al., 2009; Wilming et al., 2020) reported that gamma-band activity in visual cortex (including area MT) exhibits high coherence with the activity in (pre-) motor regions (BA 4/6). These results suggest that the relation of long-range FC and local inhibitory mechanism (MT+ GABA) support our findings that inhibitory GABA in hMT+ contributes to efficient long-range integration and coordination in distant brain areas like the prefrontal and premotor cortex.

Revised version: Another interesting finding reveal that GABAergic inhibition in hMT+ coupling with distinct functional connectivity patterns between BA 46-hMT+ and BA6-hMT+. A previous human fMRI experiment found that the positive and negative correlations between BDT and the activation of frontal regions appeared at different reasoning phases (validation or integration phases during reasoning)(Fangmeier et al., 2006). On the one hand, a monkey electrophysiological experiment reported the delayed modulation from PFC (especially in DLPFC (BA 46))) to area MT during a visual motion task(Zaksas & Pasternak, 2006). Computational models converged with empirical data of awake monkey experiments slowing temporal modulation from PFC to MT/MST(Donner et al., 2009; Siegel et al., 2015; Wang, 2002; Wimmer et al., 2015). On the other hand, human MEG studies(Donner et al., 2009; Wilming et al., 2020) reported that the gamma-band activity in the visual cortex (including area MT) exhibited high coherence with the activity in (pre-) motor regions (BA 4/6). These results suggest that the relation of long-range FC and local inhibitory mechanism (hMT+ GABA) support our findings that inhibition in hMT+ contributes to efficient long-range integration and coordination in distant brain areas like the prefrontal and premotor cortex.

Earlier version: How does hMT+ functions as a core region for 3D visuo-spatial intelligence. rather than as simple input module? The results in Figure 5a, c showed that the overlap brain regions from the analyses of connectivity-BDT-GABA/connectivity-SI-GABA are the hMT+ - BA 46. This overlap couples with local visual suppression (SI) and consequently plays an important role in 3D visuo-spatial intelligence. The direction discrimination task in this work (the visual motion paradigm of center-surround antagonism) was previously considered a mainly local function of hMT+(Melnick et al., 2013; Tadin, 2015; Tadin et al., 2003). However, our results with connectivity-SI analyses revealed that both local (FC within BA 18) and global brain connectivity (FC between hMT+ and frontal regions) contribute to SI (Table supplement 3). In human psychophysical experiments(Melnick et al., 2013; Tadin et al., 2003) the brief stimulus duration (∼100 ms) in motion discrimination precludes most top-down attentional feedback effects(Wang, 2002; Zaksas & Pasternak, 2006), while, attention, which predicted the performance of the motion discrimination task, was sustained throughout the stimulus intervals(Siegel et al., 2015). Furthermore, animal experiments have revealed that the local circuits in the visual cortex combining with top-down modulation and intracortical horizontal connection mediate the visual-spatial suppression (Angelucci et al., 2002; Keller et al., 2020; Li et al., 2019; Zhang et al., 2014).

Revised version: How does hMT+ assemble into the MD system as an intellectual hub rather than a simple input module? The results in Figure 5a, b showed that the overlap brain regions from the analyses of connectivity-BDT-GABA/connectivity-SI-GABA are the hMT+ - BA 46. This overlap couples with local visual suppression (SI) and consequently plays an important role in intelligence (BDT). The direction discrimination task in this work (the visual motion paradigm of center-surround antagonism) was previously considered a mainly local function of hMT+(Melnick et al., 2013; Tadin, 2015; Tadin et al., 2003). However, our results with connectivity-SI analyses revealed that both local (FC within BA 18) and global brain connectivity (FC between hMT+ and frontal regions) contribute to SI (Table supplement 3). In human psychophysical experiments(Melnick et al., 2013; Tadin et al., 2003) the brief stimulus duration (∼100 ms) in motion discrimination precludes most top-down attentional effects(Wang, 2002; Zaksas & Pasternak, 2006), while, attention, which predicted the performance of the motion discrimination task, was sustained throughout the stimulus intervals(Siegel et al., 2015). Furthermore, animal experiments have revealed that the local circuits in the visual cortex combining with top-down modulation and intracortical horizontal connection mediate the visual-spatial suppression(Angelucci et al., 2002; Keller et al., 2020; Li et al., 2019; Zhang et al., 2014).

Our results (shown in Figure 5a, b, right) present the intrinsic binding of local GABAergic inhibition in hMT+, which suppressing redundancy of visual motion processing (SI), and the activity of brain connectivity between hMT+ and frontal regions. These individually inherent traits may contribute to the individual difference in 3D visuo-spatial ability (Figure 5a, b, left). A candidate divisive normalization model(Carandini & Heeger, 2011; Reynolds & Heeger, 2009) can explain how such reverberation affects the process of suppressing the irrelevant information, from perception to intelligence(Melnick et al., 2013; Tadin, 2015). We summarize a framework (Figure 6) to indicate and visualize our findings.

Earlier version: Recently, Duncan et al. demonstrated coding of general fluid intelligence (gF) in distributed regions, defining them as part of multi-demand (MD) systems (Assem et al., 2020; Duncan et al., 2020). The MD system encompasses a range of cognitive domains, including working memory, mathematics, language, and relational reasoning. According to Melnick et al. (2013), motion surround suppression (SI) and time thresholds for small and large gratings, which reflect hMT+ functionality, are correlated with Verbal Comprehension, Perceptual Reasoning, Working Memory, and Processing Speed indicators. Additionally, Fedorenko et al. identified multi-demand activation regions around the occipito-temporal areas, potentially overlapping with hMT+ (Fedorenko et al., 2013). As a key region in the representation of sensory flows (including optic, tactile, and auditory flows)(Fetsch et al., 2011; Gu et al., 2006), hMT+ shows potential to be central to the MD system. Future research could focus on multi- task paradigms to further investigate the mechanisms of hMT+ and its relationship with broader cognitive functions.

Revised version: Recently, Duncan et al. demonstrated coding of general fluid intelligence (gF) in distributed regions, defining them as part of multi-demand (MD) systems (Assem et al., 2020; Duncan et al., 2020). The MD system encompasses a range of cognitive domains, including working memory, mathematics, language, and relational reasoning. According to Melnick et al. (2013), motion surround suppression (SI) and time thresholds for small and large gratings, which reflect hMT+ functionality, are correlated with Verbal Comprehension, Perceptual Reasoning, Working Memory, and Processing Speed indicators. Additionally, Fedorenko et al. identified multi-demand activation regions around the occipito-temporal areas, potentially overlapping with hMT+ (Fedorenko et al., 2013). As a key region in the representation of sensory flows (including optic, and auditory flows)(Fetsch et al., 2011; Gu et al., 2006), hMT+ shows potential to be central to the MD system. Future research could focus on multi-task paradigms to further investigate the mechanisms of hMT+ and its relationship with broader cognitive functions.

Earlier version: Together, this study offers a comprehensive insight into how the information exchange and integration between the sensory cortex (hMT+) and cognition core of BA 46, coupling with the hMT+ GABA, can predict the performance of 3D visuo-spatial ability (BDT). Our results provide direct evidence that a sensory cortex area (hMT+), its GABA biochemistry, functional connectivity, and cognition behavior levels, can assemble into complex cognition as an 3D visuo-spatial intellectual hub.

Revised version: Together, this study offers a comprehensive insight into how the information exchange and integration between the sensory cortex (hMT+) and cognition core of BA 46, coupling with the hMT+ GABA, can predict the performance of 3D visuo-spatial ability (BDT). Our results provide direct evidence that a sensory cortex area (hMT+), its GABA biochemistry, functional connectivity, and cognition behavior levels, can assemble into complex cognition as an intellectual hub.

Materials and Methods

Subjects

Earlier version: Thirty-six healthy subjects (18 female, mean age: 23.6 years±2.1, range: 20 to 29 years) participated in this study, they were recruited from Zhejiang University. All subjects had normal or corrected-to-normal vision. In addition, they reported no psychotropic medication use, no illicit drug use within the past month, no alcohol use within 3 days prior to scanning, and right-handed. This experiment was approved by the Ethics Review Committee of Zhejiang University and conducted in accordance with the Helsinki Declaration. All participants signed informed consent forms prior to the start of the study and were compensated for their time. All subjects participated in the motion spatial suppression psychophysical, resting-state fMRI and MRS (MT+ and V1 regions, in random sequence) experiments, but only part of the MRS data (31/36 in hMT+ region and 28/36 in V1 region) survived quality control (see the part of MRS date processing). The sample size is determined by the statistic requirement (30 sample for Person correlation statistical analysis).

Revised version: Thirty-six healthy subjects (18 female, mean age: 23.6 years±2.1, range: 20 to 29 years) participated in this study, they were recruited from Zhejiang University. All subjects had normal or corrected-to-normal vision. In addition, they reported no psychotropic medication use, no illicit drug use within the past month, no alcohol use within 3 days prior to scanning, and right-handed. This experiment was approved by the Ethics Review Committee of Zhejiang University and conducted in accordance with the Helsinki Declaration. All participants signed informed consent forms prior to the start of the study and were compensated for their time. All subjects participated in the motion spatial suppression psychophysical, resting-state fMRI and MRS (hMT+ and EVC (primarily in V1) regions, in random sequence) experiments, but only part of the MRS data (31/36 in hMT+ region and 28/36 in EVC (primarily in V1) region) survived quality control (see the part of MRS date processing). The sample size is determined by the statistic requirement (30 sample for Person correlation statistical analysis).

Motion surrounding suppression measurement

All stimuli were generated using Matlab (MathWorks, Natick, MA) with Psychophysics Toolbox(Brainard, 1997), and were shown on a linearized monitor (1920×1080 resolution, 100-Hz refresh rate, Cambridge Research System, Kent, UK). The viewing distance was 72 cm from the screen, with the head stabilized by a chinrest. Stimuli were drawn against a gray (56 cd per m-2) background.

A schematic of the stimuli and trial sequences is shown in our recent study(Song et al., 2021). The stimulus was a vertical drifting sinusoidal grating (contrast, 50%; spatial frequency, 1 cycle/°; speed, 4°/s) of either small (diameter of 2°) or large (diameter of 10°) size. The edge of the grating was blurred with a raised cosine function (width, 0.3°). A cross was presented in the center of the screen at the beginning of each trial for 500ms, and participants were instructed to fixate at the cross and to keep fixating at the cross throughout the trial. In each trial, a grating of either large or small size was randomly presented at the center of the screen. The grating drifted either leftward or rightward, and participants were asked to judge the perceived moving direction by a key press. Response time was not limited. The grating was ramped on and off with a Gaussian temporal envelope, and the grating duration was defined as 1 SD of the Gaussian function. The duration was adaptively adjusted in each trial, and duration thresholds were estimated by a staircase procedure. Thresholds for large and small gratings were obtained from a 160-trial block that contained four interleaved 3-down/1-up staircases. For each participant, we computed the correct rate for different stimulus durations separately for each stimulus size. These values were then fitted to a cumulative Gaussian function, and the duration threshold corresponding to the 75% correct point on the psychometric function was estimated for each stimulus size.

Stimulus demonstration and practice trials were presented before the first run. Auditory feedback was provided for each wrong response. To quantify the spatial suppression strength, we calculated the spatial suppression index (SI), defined as the difference of log10 thresholds for large versus small stimuli(Schallmo et al., 2018; Tadin et al., 2003):

Block design task measurement

The block design task was administered in accordance with the WAIS-IV manual(Wechsler, 2008). Specifically, participants were asked to rebuild the figural pattern within a specified time limit using a set of red and white blocks. The time limits were set as 30 s to 120 s according to different levels of difficulty. The patters were presented in ascending order of difficulty, and the test stopped if two consecutive patterns were not constructed in the allotted time. The score was determined by the accomplishment of the pattern and the time taken. A time bonus was awarded for rapid performance in the last six patterns. The score ranges between 0 and 66 points, with higher scores indicating better perceptual reasoning.

MR experimental procedure

Earlier version: MR experiments were performed in a 7T whole body MR system (Siemens Healthcare, Erlangen, Germany) with a Nova Medical 32 channel array head coil. Sessions included resting-state functional MRI, fMRI localizer scan, structural image scanning, and MRS scan. Resting-state scans were acquired with 1.5-mm isotropic resolution (transverse orientation, TR/TE = 2000/20.6 ms, 160 volumes, slice number = 90, flip angle = 70°, eyes closed). Structural images were acquired using a MP2RAGE sequence (TR/ TI1/ TI2 = 5000/901/3200 ms) with 0.7-mm isotropic resolution. MRS data were collected within two regions (MT+ and V1) for each subject, and we divided them into two sessions to avoid discomfort caused by long scanning. The order of MRS VOIs (MT+ and V1) in the two sessions was counterbalanced across participants. Interval between two sessions was used for block design and motion discrimination tasks. One session included fMRI localizer scan, structural image scanning, and MRS scan for the hMT+ region; the other session included structural image scan, and MRS scan for the V1 region. Spectroscopy data were acquired using a 1H-MRS single-voxel short-TE STEAM (Stimulated Echo Acquisition Mode) sequence(Frahm et al., 1989) (TE/TM/TR = 6/32/7100ms) with 4096 sampling points, 4-kHz bandwidth, 16 averages, 8 repetitions, 20×20×20 mm3 VOI size, and VAPOR (variable power and optimized relaxation delays) water suppression(Tkáč et al., 1999). Prior to acquisition, first- and second-order shims were adjusted using FASTMAP (fast, automatic shimming technique by mapping along projections(Gruetter, 1993). Two non-suppressed water spectra were also acquired: one for phase and eddy current correction (only RF pulse, 4 averages) and another for metabolite quantification (VAPOR none, 4 averages). Voxels were positioned based on anatomical landmarks using a structural image scan collected in the same session, while avoiding contamination by CSF, bone, and fat. The hMT+ VOIs were placed in the ventrolateral occipital lobe, which was based on anatomical landmarks(Dumoulin et al., 2000; Schallmo et al., 2018). We did not distinguish between the middle temporal (MT) and medial superior temporal (MST) areas in these hMT+ VOIs(Huk et al., 2002). For 14 subjects, we also functionally identified hMT+ as a check on the placement of the VOI. A protocol was used with a drifting grating (15% contrast) alternated with a static grating across blocks (10 s block duration, 160 TRs total). Using fMRI BOLD signals, these localizer data were processed online to identify the hMT+ voxels in the lateral occipital cortex, which responded more strongly to moving vs. static gratings. In addition, we only used the left hMT+ as the target region to scan, which was motivated by studies showing that left hMT+ was more effective at causing perceptual effects(Tadin et al., 2011). For V1 region, the VOI was positioned on each subject’s calcarine sulcus on the left side(Tadin et al., 2011) based on anatomical landmarks(Boucard et al., 2007; Dumoulin et al., 2000)

Revised version: MR experiments were performed in a 7T whole body MR system (Siemens Healthcare, Erlangen, Germany) with a Nova Medical 32 channel array head coil. Sessions included resting-state functional MRI, fMRI localizer scan, structural image scanning, and MRS scan. Resting-state scans were acquired with 1.5-mm isotropic resolution (transverse orientation, TR/TE = 2000/20.6 ms, 160 volumes, slice number = 90, flip angle = 70°, eyes closed). Structural images were acquired using a MP2RAGE sequence (TR/ TI1/ TI2 = 5000/901/3200 ms) with 0.7-mm isotropic resolution. MRS data were collected within two regions (hMT+ and EVC (primarily in V1)) for each subject, and we divided them into two sessions to avoid discomfort caused by long scanning. The order of MRS VOIs (hMT+ and EVC (primarily in V1)) in the two sessions was counterbalanced across participants. Interval between two sessions was used for block design and motion discrimination tasks. One session included fMRI localizer scan, structural image scanning, and MRS scan for the hMT+ region; the other session included structural image scan, and MRS scan for the EVC (primarily in V1) region. Spectroscopy data were acquired using a 1H-MRS single-voxel short-TE STEAM (Stimulated Echo Acquisition Mode) sequence(Frahm et al., 1989) (TE/TM/TR = 6/32/7100ms) with 4096 sampling points, 4-kHz bandwidth, 16 averages, 8 repetitions, 20×20×20 mm3 VOI size, and VAPOR (variable power and optimized relaxation delays) water suppression(Tkáč et al., 1999). Prior to acquisition, first- and second-order shims were adjusted using FASTMAP (fast, automatic shimming technique by mapping along projections(Gruetter, 1993). Two non-suppressed water spectra were also acquired: one for phase and eddy current correction (only RF pulse, 4 averages) and another for metabolite quantification (VAPOR none, 4 averages). Voxels were positioned based on anatomical landmarks using a structural image scan collected in the same session, while avoiding contamination by CSF, bone, and fat. The hMT+ VOIs were placed in the ventrolateral occipital lobe, which was based on anatomical landmarks(Dumoulin et al., 2000; Schallmo et al., 2018). We did not distinguish between the middle temporal (MT) and medial superior temporal (MST) areas in these hMT+ VOIs(Huk et al., 2002). For 14 subjects, we also functionally identified hMT+ as a check on the placement of the VOI. A protocol was used with a drifting grating (15% contrast) alternated with a static grating across blocks (10 s block duration, 160 TRs total). Using fMRI BOLD signals, these localizer data were processed online to identify the hMT+ voxels in the lateral occipital cortex, which responded more strongly to moving vs. static gratings. In addition, we only used the left hMT+ as the target region to scan, which was motivated by studies showing that left hMT+ was more effective at causing perceptual effects(Tadin et al., 2011). For EVC (primarily in V1) region, the VOI was positioned on each subject’s calcarine sulcus on the left side(Tadin et al., 2011) based on anatomical landmarks(Boucard et al., 2007; Dumoulin et al., 2000)

MRS data processing

Spectroscopy data were preprocessed and quantified using magnetic resonance signal processing and analysis, https://www.cmrr.umn.edu/ downloads/mrspa/), which runs under MATLAB and invokes the interface of the LCModel (Version 6.3-1L) (Chen et al., 2019). First, we used the non-suppressed water spectra to perform eddy current correction and frequency/phase correction. Second, we checked the quality of each FID (16 averages) visually and removed those with obviously poor quality. Third, the absolute concentrations of each metabolite were quantitatively estimated via the Water- Scaling method. For partial-volume correction, the tissue water content was computed as follows(Ernst et al., 1993): where fgm, fwm, and fcsf were the GM/WM/CSF volume fraction in MRS VOI and we used FAST (fMRI’s automated segmentation tool, part of the FSL toolbox)(Zhang et al., 2001) to segment the three tissue compartments from the T1-weighted structural brain images. For water T2 correction, we set water T2 as 47ms (Marjańska et al., 2012). Our concentrations were mM per kg wet weight. Furthermore, LCModel analysis was performed on all spectra within the chemical shift range of 0.2 to 4.0 ppm.

Poor spectral quality was established by a Cramer-Rao Lower Bound (CRLB) of more than 20%(Cavassila et al., 2001), and some data were excluded from further analysis. The details were described in our recently paper(Song et al., 2021).

Rs-fMRI data processing and analysis

Resting-state functional image was analyzed in the Data Processing and Analysis for Brain Imaging DPABI toolbox(Yan et al., 2016) based on SPM 12 (http://www.fil.ion.ucl.ac.uk/spm/). The preprocessing steps included discard of the first five volumes, slice timing, realignment to the 90th slice, coregistration of each subject’s T1-weighted anatomical and functional images, segmentation of the anatomical images into six types of tissues using DARTEL, linear detrend, regressing nuisance variables (including realignment Friston 24-parameter, global signal, white matter and CSF signal) (Friston et al., 1996), normalization to the standard Montreal Neurological Institute (MNI) space with the voxel size of 1.5× 1.5× 1.5 mm3 using DARTEL, spatial smoothing with a Gaussian kernel of 3 mm full-width-half-maximum (FWHM), and band-pass filtering with Standard frequency band (SFB, 0.01–0.1 Hz). Spherical ROI with a radius of 6mm was placed in left MT. The coordinate for left MT (-46, -72, -4, in MNI space) was obtained by our localizer fMRI experiment. We calculated the seed-to-voxel whole brain FC map for each subject. All the FC values were Fisher-Z-transformed.

We did a similar connectivity-behavior analysis to a previous study (Song et al., 2008). First, we computed the Pearson’s correlation coefficient between BDT scores and the FC values across subjects in a voxel-based way. Then, to evaluate the significance, we transformed the r-value into t-value , where df denotes to the degrees of freedom, and r is the Pearson’s correlation coefficient between BDT scores and the FC values. Here, df was equal to 27. The brain regions in which the FC values to the seed region was significantly correlated with the BDT scores were obtained with a threshold of P < 0.005 for regions of a priori (|t(27) | ≥ 3.057, and adjacent cluster size ≥ 23 voxels (AlphaSim corrected)), and P < 0.01 for whole-brain analyses (|t(27) | ≥ 2.771 and adjacent cluster size ≥ 37 voxels (AlphaSim corrected).

Statistical Analysis

PROCESS version 3.4, a toolbox in SPSS, was used to examine the mediation model. There are some prerequisites for mediation analysis: the independent variable should be a significant predictor of the mediator, and the mediator should be a significant predictor of the dependent variable.

SPSS 20 (IBM, USA) was used to conduct all the remaining statistical analysis in the study. We evaluated the correlation of variables (GABA, Glu, SI, BDT) using Pearson’s correlation analysis. Differences or correlations were considered statistically significant if P < 0.05. Significances with multiple comparisons were tested with false discovery rate (FDR) correction. The effect of age on intelligence was controlled for by using partial correlation in the correlation analysis and was taken as a covariate in the serial mediation model analysis.

Data availability

Earlier version: Source data are provided with this paper and have been archived at GitHub and could be downloaded with reasonable request. https://github.com/Yrehearsal/GABA-_hMT-_Intelligence.

Revised version: Source data are provided with this paper and have been archived at zenodo and could be downloaded with reasonable request, https://doi.org/10.5281/zenodo.12789781

Code availability

Added in the revised version: This code has been uploaded to the GitHub: https://github.com/Yrehearsal/GABA-_hMT-_Intelligence

Removed in the revised version: This paper does not report original code.

Removed in the revised version: Supporting information